## 문제1
rain 변수에는 시애틀 지역의 2014년 강우량에 대한 정보가 numpy array 형태로 저장되어 있다. 데이터 내에는 비가 내린날과 그렇지 않은날이 모두 존재한다. 다음 요건을 보고, 각 변수가 요건을 만족할수 있도록 만드시오
* 변수 result1에 비가 내린날의 평균 강우량을 저장하시오
* 변수 result2에 비가 내리지 않은 날의 갯수를 저장하시오
* 변수 result3에 비가 내린 날의 강우량의 최대값과 최소값의 차이를 계산하여 저장하시오

<!-- grade_id: problem1, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np

## 다음 코드를 건드리지 마시오
rain=pd.read_csv('https://raw.githubusercontent.com/ralbu85/DataProcessing1_2022S/master/Seattle2014.csv')['PRCP'].values
print(rain)
# YOUR CODE HERE
raise NotImplementedError()
print(result1, result2, result3)
```

<!-- grade_id: problem1_test, points: 4, grade: true, solution: false -->
```python
from nose.tools import assert_equal
print('success')
```

## 문제2
points라는 변수에 (100,2) 데이터셋이 무작위로 생성되었다. 여기서 각 행은 2차원 공간에 찍힌 점을 의미한다. target에 저장되어 있는 [0.5,0.5]과 가장 가까운 10개의 점을 찾고 이들의 평균 거리를 계산하여 result 변수에 저장하시오 

<!-- grade_id: problem2, points: 0, grade: false, solution: true -->
```python
import numpy as np
rand=np.random.RandomState(42)
points=rand.rand(100,2)
target = np.array([0.5,0.5])
# YOUR CODE HERE
raise NotImplementedError()
print(result)
```

<!-- grade_id: problem2_answer, points: 4, grade: true, solution: false -->
```python
import numpy as np
print('success')
```

## 문제3
손글씨 이미지를 나타내기 위한 데이터셋을 data라는 변수에 불러왔다. 이 데이터는 현재 1797개의 이미지로 구성되어 있으며 각 이미지는 64개의 픽셀로 표현되고 있다.즉 데이터가 (1797,64)로 표현된다. result 변수를 만들고 다음 요건을 만족할 수 있도록 하시오. 
* 요건1: 현재 데이터의 차원은 (1797,64)인데 이를 (1797,8,8)로 바꾸시오
* 요건2: 현재 데이터의 값이 0에서 16까지의 값으로 표현되고 있는데 이를 0~1 사이의 값을 가질 수 있도록 변형하시오



<!-- grade_id: problem3, points: 0, grade: false, solution: true -->
```python
from sklearn.datasets import load_digits
## 다음 코드를 건드리지 마시오
data = load_digits()['data']
print(data, data.shape)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem3_answer, points: 4, grade: true, solution: false -->
```python
print('success')
```

## 문제4

손글씨 이미지를 나타내기 위한 데이터셋을 data라는 변수에 불러왔다. 이 데이터는 현재 1797개의 이미지로 구성되어 있으며 각 이미지는 64개의 픽셀로 표현되고 있다.즉 데이터가 (1797,64)로 표현된다. point라는 변수는 전체 데이터 중 14번째 데이터를 가리키고 있는데, point 자기 자신을 제외한 가장 가까운 10개의 데이터를 data에서 검색하고, 그 위치를 argsort를 이용하여 구해보시오. 계산 결과는 result 변수에 저장한다

<!-- grade_id: problem4, points: 0, grade: false, solution: true -->
```python
from sklearn.datasets import load_digits
## 다음 코드를 건드리지 마시오
data = load_digits()['data']
point = data[13]
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem4_test, points: 4, grade: true, solution: false -->
```python
print('success')
```

## 문제5
data라는 변수에 (100,5) 데이터셋이 무작위로 생성되었다. 각 행별로(100개의) 합을 구하고, 이 합에 따라 data의 행을 재정렬하여 result에 저장하시오. result의 첫번째 행에는 5개 숫자의 합이 가장 작은 행이 오게 될 것이며, result의 마지막 행에는 5개의 합이 가장 컸던 행이 오게 될 것이다.

<!-- grade_id: problem5, points: 0, grade: false, solution: true -->
```python
import numpy as np
rand=np.random.RandomState(24)
data=rand.rand(100,5)
# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: problem5_test, points: 4, grade: true, solution: false -->
```python
import numpy as np
print('success')
```

## 문제6
data에는 1949년 1월부터 1960년 12월까지 월별 항공기 이용 여객수가 기록되어 있다. 현재 1차원인 이 데이터를 변경하여 2차원 데이터로 만들고 이를 result에 저장하고 싶다. result는 다음과 같은 규칙으로 구성되어있다. 이를 만족하도록 2차원 array를 만들고 result에 저장하시오
* 1행, data의 1번째값~5번째 값 [112,118,132,129,121]
* 2행, data의 2번째값~6번째 값 [118,132,129,121,135]
* ...
* 마지막행, data의 마지막 5개 값 [606,508,461,390,432]

<!-- grade_id: problem6, points: 0, grade: false, solution: true -->
```python
import pandas as pd
data=pd.read_csv('https://raw.githubusercontent.com/jbrownlee/Datasets/master/airline-passengers.csv')['Passengers'].values
print(data)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem6_test, points: 4, grade: true, solution: false -->
```python
import numpy as np
print('success')
```

## 문제7

data에는 1949년 1월부터 1960년 12월까지 월별 항공기 이용 여객수가 순서대로 기록되어 있다. 12개월 단위로 데이터를 나누어 매해 평균 이용객을 계산하여 result 변수에 저장하시오

<!-- grade_id: problem7, points: 0, grade: false, solution: true -->
```python
import pandas as pd
data=pd.read_csv('https://raw.githubusercontent.com/jbrownlee/Datasets/master/airline-passengers.csv')['Passengers'].values
print(data)
# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: problem7_test, points: 4, grade: true, solution: false -->
```python
import numpy as np
print('success')
```